<problem>
<startouttext/>
<endouttext/>

<script type="loncapa/python">

def test_add_to_ten(expect,ans):
  a1=float(ans[0])
  a2=float(ans[1])
  return (a1+a2)==10


</script>

<p>Enter two integers which sum to 10: </p>
<customresponse cfn="test_add_to_ten">
        <textline size="40"/><br/>
        <textline size="40"/>
</customresponse>

    <solution></solution>
       
</problem>