<?xml version="1.0"?>
<problem showanswer="closed" rerandomize="never" weight="10">
  <text>
    <p>
      <h4>Example "option" problem</h4>
    </p>
    <p>
Where is the earth? </p>
    <p>
      <optionresponse>
        <optioninput options="('up','down')" correct="down"/>
      </optionresponse>
    </p>
    <p>
      <h4>Example "symbolic" problem</h4>
    </p>
    <p>
What is Einstein's equation for the energy equivalent of a mass [mathjaxinline]m[/mathjaxinline]? </p>
    <p>
      <symbolicresponse expect="m*c^2">
        <textline size="90" correct_answer="m*c^2" math="1"/>
      </symbolicresponse>
    </p>
    <p>
      <h4>Example "numerical" problem</h4>
    </p>
    <p>
Estimate the energy savings (in J/y) if all the people ([mathjaxinline]3\times 10^8[/mathjaxinline]) in the U.&#xA0;S. switched from U.&#xA0;S. code to low flow shower heads. </p>
    <p>
      <p style="display:inline">Energy saved = </p>
      <numericalresponse inline="1" answer="0.52">
        <textline inline="1">
          <responseparam type="tolerance" default="0.02"/>
        </textline>
      </numericalresponse>
      <p style="display:inline">&#xA0;EJ/year</p>
    </p>
    <p>
      <h4>Example "multiple choice" problem</h4>
    </p>
    <p>
What color is a banana? </p>
    <p>
      <choiceresponse>
        <checkboxgroup>
          <choice correct="false" name="1">
            <text>Red</text>
          </choice>
          <choice correct="false" name="2">
            <text>Green</text>
          </choice>
          <choice correct="true" name="3">
            <text>Yellow</text>
          </choice>
          <choice correct="false" name="4">
            <text>Blue</text>
          </choice>
        </checkboxgroup>
      </choiceresponse>
    </p>
    <p>
      <h4>Example "string response" problem</h4>
    </p>
    <p>
In what U.S. state is Detroit located? </p>
    <p>
      <stringresponse answer="Michigan">
        <textline/>
      </stringresponse>
    </p>
    <p>
An explanation of the answer can be provided by using the edXsolution macro. Click on "Show Answer" to see the solution. </p>
    <p>
      <solution>
        <font color="blue">Answer: </font>
        <font color="blue">Detroit is near Canada, but it is actually in the United States. </font>
      </solution>
    </p>
    <p>
      <h4>Example "custom response" problem</h4>
    </p>
    <p>
This problem demonstrates the use of a custom python script used for checking the answer. </p>
    <script type="text/python" system_path="python_lib">
def sumtest(expect,ans):
    (a1,a2) = map(float,eval(ans))
    return (a1+a2)==10
</script>
    <p>
Enter a python list of two numbers which sum to 10, eg [9,1]: </p>
    <p>
      <customresponse cfn="sumtest" expect="[1,9]">
        <textline correct_answer="[1,9]"/>
      </customresponse>
    </p>
    <p>
      <h4>Example image</h4>
    </p>
    <p>
Include image by using the edXxml macro: </p>
    <p>
      <img src="http://web.mit.edu/jorloff/www/cube-tet/cube_tet.gif"/>
    </p>
    <p>
      <h4>Example show/hide explanation</h4>
    </p>
    <p>
Extra explanations can be tucked away behind a "showhide" toggle flag: </p>
    <p>
      <table class="wikitable collapsible collapsed">
        <tbody>
          <tr>
            <th> More explanation [<a href="javascript:$('#sh1').toggle()" id="sh1l">show</a>]</th>
          </tr>
          <tr id="sh1" style="display:none">
            <td>
              <p>
This is a hidden explanation. It can contain equations: [mathjaxinline]\alpha = \frac{2}{\sqrt {1+\gamma }}[/mathjaxinline]  </p>
              <p>
This is some text after the showhide example. </p>
            </td>
          </tr>
        </tbody>
      </table>
    </p>
  </text>
</problem>
